# Final Project: Your Choice

You have made it to the final week of the semester! 

### Logistics

**Rules**: You are required to choose your own topic and complete your own work, with the help of the instructor if necessary. If you are using outside sources such as example code generated by AI, etc. please cite the source or the prompt that you used to generate the code. 

**Programming Environment**: You may choose to complete your work on the same notebook server as used for assignments and labs. You can add your work to this notebook and submit the completed the ```.ipynb``` and ```.html``` files in Canvas as you did for the assignments. If you need to use additional libraries that are not available through ```import```, you can use the shell command to install them. For excample, if you would like to use the ```pytorch``` library for training neural nets, you can do: 

```!pip3 install torch```

Alternatively, you can complete all programming on your own device. Simply download this notebook and run it on your device. 

**Support**: You are not alone! Come to office hours. If you're ever feeling overwhelmed or don't know how to make progress, email for help.

### Coverage
The final project requires you to explore and analyze your own data set by leveraging the resources and ideas you learned during this semester. You can elect to collect the data yourself, use something from the class repository, or from other third-party sources. If you use online sources, you must cite the source of your data, and your statistical analysis must be original. Your project should cover each of the following areas:

**Descriptive statistics**: graphic and numeric summaries of your data

**Statistical Inference**: show examples of statistical inference by apply the techniques from this class. For example, based on the student data, we may ask:
Is there a difference in average between students who do and do not use social media? (A/B testing)
What is a 95% confidence interval for the difference between the average height of male and female students? (confidence interval by bootstrapping)
Is there evidence to show there is a correlation between the Shoe size and Height of the male students? Can you predict the shoe size for a student who is 76 inches tall? (correlation and regression)

**Classification**: show how you can predict the outcome of one of the variables by using a set of features by using a classifer (such as kNN, or other models available through ```sklearn```. Evaluate the performance of your classifier and discuss areas of further refinement.
You should ask and address at least two areas mentioned above.



```python
# Run this cell to set up the notebook, but please don't change it.
import numpy as np
import math, scipy, sklearn
import datascience

from datascience import *

# These lines set up the plotting functionality and formatting.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter("ignore")

```

### Introduction (The Dataset)

Describe your data below, including the necesary background and the meaning of the variables. If you have downloaded the data from the web, please identify its source. 



Tests for q1_0 (points: [0])
```
>>> 0 < expected_row_sum
True
```


Tests for q1_1_1 (points: [0, 0])
```
>>> type(stemmed_message) == str
True
```
```
>>> len(stemmed_message) < len('elements')
True
```


Tests for q1_1_2 (points: [0])
```
>>> type(most_stem) == str
True
```


Tests for q1_1_3 (points: [0])
```
>>> type(longest_uncut) == str
True
```


Tests for q1_1_4 (points: [0])
```
>>> type(count_single_stems) == int
True
```


Tests for q1_2_1 (points: [0])
```
>>> 0.2 < outer_space_r < 0.4
True
```


Tests for q1_2_3 (points: [0])
```
>>> type(san_francisco) == int
True
```


Tests for q2_1_1 (points: [0])
```
>>> 0 < one_distance < 0.01
True
```


Tests for q2_1_2 (points: [1, 2])
```
>>> correct_dis = 0.001406116
>>> dis = distance_two_features('clerks.', 'the godfather', 'water', 'feel')
>>> np.isclose(np.round(dis, 9), correct_dis)
True
```
```
>>> correct_dis = 0.005647119
>>> dis = distance_two_features('clerks.', 'the godfather', 'your', 'that')
>>> np.isclose(np.round(dis, 9), correct_dis)
True
```


Tests for q2_1_3 (points: [0])
```
>>> np.isclose(distance_from_python('clerks.'), 0.00079838)
True
```


Tests for q2_1_4 (points: [0, 0, 0])
```
>>> set(close_movies.labels) >= {'Genre', 'Title', 'feel', 'water'}
True
```
```
>>> close_movies.num_rows == 5
True
```
```
>>> close_movies.column('Title').item(0) != 'monty python and the holy grail'
True
```


Tests for q2_1_5 (points: [3])
```
>>> [most_common('Genre', close_movies.take(range(k))) for k in range(1, 6, 2)]
['thriller', 'thriller', 'thriller']
```


Tests for q3_0 (points: [0, 0, 0])
```
>>> 0.0 <= distance_first_to_second <= 0.1
True
```
```
>>> np.isclose(distance(make_array(1, 2), make_array(1, 2)), 0)
True
```
```
>>> np.isclose(distance(make_array(1, 2, 3), make_array(2, 4, 5)), 3)
True
```


Tests for q3_1_1 (points: [0])
```
>>> bottom_left >= 1 and bottom_left <= 5
True
```


Tests for q3_1_2 (points: [0])
```
>>> bottom_right >= 1 and bottom_right <= 5
True
```


Tests for q3_1_3 (points: [0])
```
>>> top_right >= 1 and top_right <= 5
True
```


Tests for q3_1_4 (points: [0])
```
>>> top_left >= 1 and top_left <= 5
True
```


Tests for q3_1_5 (points: [0])
```
>>> movie_genre_guess >= 1 and movie_genre_guess <= 2
True
```


Tests for q3_1_6 (points: [0, 0, 0, 1.5, 1.5])
```
>>> len(my_features) >= 10
True
```
```
>>> np.all([f in test_movies.labels for f in my_features])
True
```
```
>>> len(set(my_features)) >= 10
True
```
```
>>> help_message = "It looks like there are many movies in the training set that \n                  don't have any of your chosen words.  That will make your\n                  classifier perform very poorly in some cases.  Try choosing\n                  at least 1 common word.\n                  "
>>> train_f = train_movies.select(my_features)
>>> np.count_nonzero(train_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < len(my_features)
True
```
```
>>> help_message = "It looks like there are many movies in the test set that \n                  don't have any of your chosen words.  That will make your\n                  classifier perform very poorly in some cases.  Try choosing\n                  at least 1 common word.\n                  "
>>> test_f = test_movies.select(my_features)
>>> np.count_nonzero(test_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < 5
True
```


Tests for q3_1_8 (points: [0, 1, 1, 1])
```
>>> genre_and_distances.labels == ('Genre', 'Distance')
True
```
```
>>> genre_and_distances.num_rows == train_movies.num_rows
True
```
```
>>> print(genre_and_distances.group('Genre'))
Genre    | count
comedy   | 102
thriller | 181
```
```
>>> np.allclose(genre_and_distances.column('Distance'), sorted(fast_distances(test_my_features.row(0), train_my_features)))
True
```


Tests for q3_1_9 (points: [1, 1])
```
>>> genre_and_distances.take(np.arange(7)).group('Genre').index_by('Genre')[my_assigned_genre][0].item('count') >= 4
True
```
```
>>> my_assigned_genre_was_correct == (my_assigned_genre == 'thriller')
False
```


Tests for q3_2_1 (points: [1.5, 1.5])
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> 
>>> def check(r, k):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, g, k) == Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:k])).most_common(1)[0][0]
>>> check_5_nn = [check(i, 5) for i in np.arange(11)]
>>> all(check_5_nn)
True
```
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> 
>>> def check(r, k):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, g, k) == Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:k])).most_common(1)[0][0]
>>> check_11_nn = [check(i, 11) for i in np.arange(11)]
>>> all(check_11_nn)
True
```


Tests for q3_2_2 (points: [3])
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> r = np.where(test_movies['Title'] == 'godzilla')[0][0]
>>> t = test_my_features.row(r)
>>> godzilla_expected_genre = Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:15])).most_common(1)[0][0]
>>> godzilla_genre == godzilla_expected_genre
True
```


Tests for q3_2_3 (points: [3])
```
>>> def check(r):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, train_movies.column('Genre'), 15) == classify_feature_row(t)
>>> all([check(i) for i in np.arange(15)])
True
```


Tests for q3_3_1 (points: [0])
```
>>> 0 <= proportion_correct <= 1
True
```


Tests for q3_3_2 (points: [0, 0, 3])
```
>>> test_movie_correctness.labels == ('Title', 'Genre', 'Was correct')
True
```
```
>>> test_movie_correctness.num_rows == test_movies.num_rows
True
```
```
>>> print(test_movie_correctness.group('Genre'))
Genre    | count
comedy   | 17
thriller | 33
```


Tests for q4_1 (points: [0, 3])
```
>>> len(new_features) <= 5
True
```
```
>>> another_classifier(test_new.row(0)) in ['comedy', 'thriller']
True
```
